```python
pip install nose
```

Output:
```
Requirement already satisfied: nose in /opt/conda/lib/python3.10/site-packages (1.3.7)
Note: you may need to restart the kernel to use updated packages.
```

```python
#Perform logistic regression on diabetes dataset
```

<!-- grade_id: cell-65901a02d7b3d454, points: 0, grade: false, solution: true -->
```python
#Load diabetes dataset here and return shape of dataset
from cz_utility import dataset
# diabetes dataset
df = dataset.load('63ea33b29ac5f900139b8d07')
def returnShape(df):
    # your code here
    return df.shape
    raise NotImplementedError
```

<!-- grade_id: cell-9d4cd5b17d896c66, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(returnShape(df), (768, 9))
```

```python
df.head(10)
```

Output:
```
Pregnancies  Glucose  BloodPressure  SkinThickness  Insulin   BMI  \
0            6      148             72             35        0  33.6   
1            1       85             66             29        0  26.6   
2            8      183             64              0        0  23.3   
3            1       89             66             23       94  28.1   
4            0      137             40             35      168  43.1   
5            5      116             74              0        0  25.6   
6            3       78             50             32       88  31.0   
7           10      115              0              0        0  35.3   
8            2      197             70             45      543  30.5   
9            8      125             96              0        0   0.0   

   DiabetesPedigreeFunction  Age  Outcome  
0                     0.627   50        1  
1                     0.351   31        0  
2                     0.672   32        1  
3                     0.167   21  
```

```python
(df['Insulin']==0).sum()
```

Output:
```
374
```

<!-- grade_id: cell-97b7b8d9c24aa579, points: 0, grade: false, solution: true -->
```python
#separate input and output from dataset and remove unwanted data columns
#take x for input and y for output
def sepAndRemove(df):
    # your code here
    #The unwanted column is 'SkinThickness' because it does not affect diabetes.
    x=df.drop(['Insulin','Outcome'],axis=1)
    y=df['Outcome']
    return x,y
    raise NotImplementedError
```

<!-- grade_id: cell-f7e39f11fcd55f39, points: 2, grade: true, solution: false -->
```python
trainAndTestData = sepAndRemove(df)
assert_equal(trainAndTestData[0].shape, (768, 7))
assert_equal(trainAndTestData[1].shape, (768,))
```

<!-- grade_id: cell-05c4eb3e485c557a, points: 0, grade: false, solution: true -->
```python
#separate training and testing dataset with test_size = 0.3
from sklearn.model_selection import train_test_split
def trainTestSplit(trainAndTestData):
    # your code here
    X, y = trainAndTestData

    # Perform train-test split with test_size = 0.3 and random_state = 42
    X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=42)

    # Print shapes for verification
    print("X_train shape:", X_train.shape)
    print("X_test shape:", X_test.shape)
    print("y_train shape:", y_train.shape)
    print("y_test shape:", y_test.shape)

    return X_train, X_test, y_train, y_test
    raise NotImplementedError
```

<!-- grade_id: cell-a7a6e742dbef0d20, points: 2, grade: true, solution: false -->
```python
trainAndTestData=trainTestSplit(sepAndRemove(df))
assert_equal(trainAndTestData[0].shape,(537, 7))
assert_equal(trainAndTestData[1].shape,(231, 7))
assert_equal(trainAndTestData[2].shape,(537,))
assert_equal(trainAndTestData[3].shape,(231,))
```

Output:
```
X_train shape: (537, 7)
X_test shape: (231, 7)
y_train shape: (537,)
y_test shape: (231,)
```

<!-- grade_id: cell-6e973aeb54f8a469, points: 0, grade: false, solution: true -->
```python
#use MinMaxScaler to scale the data to a fixed range
from sklearn.preprocessing import MinMaxScaler
def scalar(trainAndTestData):
    # your code here
    # Unpack train and test data
    X_train, X_test, y_train, y_test = trainAndTestData
    
    # Instantiate MinMaxScaler
    scaler = MinMaxScaler()

    # Fit scaler on the training data and transform both training and test data
    X_train_scaled = scaler.fit_transform(X_train)
    X_test_scaled = scaler.transform(X_test)
    
    return X_train_scaled, X_test_scaled
    raise NotImplementedError
```

<!-- grade_id: cell-0bb62bf101293688, points: 2, grade: true, solution: false -->
```python
trainAndTestData=scalar(trainTestSplit(sepAndRemove(df)))
assert_equal(trainAndTestData[0].shape, (537, 7))
assert_equal(trainAndTestData[1].shape, (231, 7))

```

Output:
```
X_train shape: (537, 7)
X_test shape: (231, 7)
y_train shape: (537,)
y_test shape: (231,)
```

<!-- grade_id: cell-a9178e0adf01003e, points: 0, grade: false, solution: true -->
```python
#perform logistic regression here and return accuracy between predected output for testx and testy
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score
def logisticReg(trainAndTestData, trainy, testy):
    # your code here
    model=LogisticRegression()
    X_train,X_test=trainAndTestData
    model.fit(X_train,trainy)
    y_pred=model.predict(X_test)
    accuracy=accuracy_score(testy,y_pred)
    return accuracy
    raise NotImplementedError
```

<!-- grade_id: cell-b26830e8ae1d3a06, points: 2, grade: true, solution: false -->
```python
#there might be difference between student answer and faculty answer
trainAndTestData1 = trainTestSplit(sepAndRemove(df))
accuracyScore = logisticReg(scalar(trainTestSplit(sepAndRemove(df))), trainAndTestData1[2], trainAndTestData1[3])
assert_equal(accuracyScore>0.60, True)
```

Output:
```
X_train shape: (537, 7)
X_test shape: (231, 7)
y_train shape: (537,)
y_test shape: (231,)
X_train shape: (537, 7)
X_test shape: (231, 7)
y_train shape: (537,)
y_test shape: (231,)
```

```python
print(accuracyScore)
```

Output:
```
0.7445887445887446
```

```python

```